# CMPS 140

## Assignment 1

**DUE: Friday April 20, 2018 11:59pm**

Turn in the assignment via Canvas.

To write legible answers you will need to be familiar with both [Markdown](https://github.com/adam-p/markdown-here/wiki/Markdown-Cheatsheet) and [Latex](https://www.latex-tutorial.com/tutorials/amsmath/)

Before you turn this problem in, make sure everything runs as expected. First, restart the kernel (in the menubar, select Kernel→→Restart) and then run all cells (in the menubar, select Cell→→Run All).

Make sure you fill in any place that says "YOUR CODE HERE" or "YOUR ANSWER HERE", as well as your name below:


```python
NAME = "Kenny Luu"
STUDENT_ID = "1491158"
```

## Problem 1 

Suppose that 13 cards are selected at random from a regular deck of 52 playing cards:

a. If it is known that at least one ace has been selected, what is the probability that at least two aces have been selected?

b. If it is known that the ace of hearts has been selected, what is the probability that at least two aces have been selected?

<!-- grade_id: cell-b18eb48e93041a25, points: 1, grade: true, solution: true -->
# Answers
## A.

Let **A** be the event of getting at least 2 Aces when 13 cards are selected:

$P(A) = \dfrac{{{4}\choose{2}} \cdot {{48}\choose{11}} + {{4}\choose{3}} \cdot {{48}\choose{10}} + {{4}\choose{4}} \cdot {{48}\choose{9}}}{{52}\choose{13}}$

Let **B** be the event of getting at least one Ace when 13 cards are selected:

$P(B) = \dfrac{{{4}\choose{1}} \cdot {{48}\choose{12}} + {{4}\choose{2}} \cdot {{48}\choose{11}} + {{4}\choose{3}} \cdot {{48}\choose{10}} + {{4}\choose{4}} \cdot {{48}\choose{9}}}{{52}\choose{13}}$

So the probability of A given B is

$P(A|B) = \dfrac{P(A\cap B)}{P(B)}$

Since $P(A\cap B)$ is the same as $P(A)$ we get:

$P(A|B) = \dfrac{P(A)}{P(B)}$ which is equal to: $.365$

## B.

Let A be the event of getting at least 2 Aces when 13 cards are selected:

$P(A) = \dfrac{{{4}\choose{2}} \cdot {{48}\choose{11}} + {{4}\choose{3}} \cdot {{48}\choose{10}} + {{4}\choose{4}} \cdot {{48}\choose{9}}}{{52}\choose{13}}$

Let B be the event that the Ace of hearts is chosen when 13 cards are selected

$P(B) = \dfrac{{51}\choose{12}}{{52}\choose{13}}$




## Problem 2


Suppose that a person's score X on a mathematics aptitude test is a number between 0 and 1, and that a person's score on a music aptitude test is also a number between 0 and 1. Supposed further that in the population of all college students, the scores X and Y are distributed according to the following joint p.d.f.:

$f(x, y) = \frac{2}{5}(2x+3y),$ for $0 \leq x \leq 1$ and $0 \leq y \leq 1$

a. What proportion of college students obtain a score greater than 0.8 on the mathematics test?

b. If a student's score on the music test is 0.3, what is the probability that his/her score on the mathematics test will be greater than 0.8?

c. If a student's score on the mathematics test is 0.3, what is the probability that his/her score on the music test will be greater than 0.8?

<!-- grade_id: cell-2f061f10915044a8, points: 1, grade: true, solution: true -->
[YOUR ANSWER HERE]

## Problem 3

Write out the joint probability distribution described by the following Bayesian Network:

![bayes_net](https://docs.google.com/drawings/d/e/2PACX-1vRjJ382tPK3UrxydjCmwWmMjOi-8T1wQN6hBBJhE7u3-creWY7KR-K224MPTU5IzKXtX4ILxviRUMAw/pub?w=720&h=528)

<!-- grade_id: cell-e702292884b97e63, points: 1, grade: true, solution: true -->
[YOUR ANSWER HERE]

## Problem 4

Given the dataset below, use information gain to calculate the first attribute the decision tree should split on. Feel free to use the cell below to explore the dataset and facilitate with the calculations. You should show your work using LaTeX 2 cells down.

```python
import numpy as np
from sklearn.datasets import load_iris

data_iris, labels_iris = load_iris(return_X_y=True)

print('There are {} samples with {} attributes.'.format(*data_iris.shape))
print('The labels consist of {} classes'.format(len(np.unique(labels_iris))))
```

<!-- grade_id: cell-9d85836f00c56d8e, points: 0, grade: true, solution: true -->
[YOUR ANSWER HERE]

## Problem 5

Given the dataset below, implement the logistic regression algorithm, optimizing the parameters using gradient descent and squared error as the loss function. See 18.6.4 from the textbook for details.

```python
import numpy as np
import matplotlib.pyplot as plt
from sklearn.datasets.samples_generator import make_moons

np.random.seed(33)
data, labels = make_moons(n_samples=100, noise=0.1)
colors = ['r' if y else 'b' for y in labels]
print(data.shape, labels.shape)
plt.scatter(data[:,0], data[:,1], c=colors)
plt.show()
```

<!-- grade_id: cell-14e8f247a6e4f45c, points: 0, grade: false, solution: true -->
```python
def logistic(x, w):
    # YOUR CODE HERE
    raise NotImplementedError()
    
def sigmoid(z):
    # YOUR CODE HERE
    raise NotImplementedError()

def sigmoid_derivative(z):
    # YOUR CODE HERE
    raise NotImplementedError()
    
def squared_error(y, h):
    # YOUR CODE HERE
    raise NotImplementedError()

def squared_error_derivative(y, h):
    # YOUR CODE HERE
    raise NotImplementedError()
    
def accuracy(y, h):
    # YOUR CODE HERE
    raise NotImplementedError()
    
def logistic_regression(x, y, loss_func, loss_func_derivative, learning_rate, num_steps=10):
    # start with intial parameters w_i = 1
    w = np.ones(3)
    # include a bias term
    x = np.pad(x, [[0,0], [1,0]], mode='constant')

    print('Intial Accuracy:{}%'.format(accuracy(y, np.round(logistic(x, w)))))
    
    for step in range(num_steps):
        output = logistic(x, w)
        update = np.dot(loss_func(y, output) * loss_func_derivative(y, output), x)
        w = w + learning_rate * update
        print('Step {} Accuracy:{}%'.format(step+1,accuracy(y, np.round(logistic(x, w)))))
    return w
    

ws = logistic_regression(data, labels, squared_error, squared_error_derivative, 0.5)
yh = np.round(logistic(np.pad(data, [[0,0], [1,0]], mode='constant'), ws))
colors = ['g' if _yh==_y else 'r' for _yh, _y in zip(yh, labels.astype(np.int))]
plt.title('Classification Results')
plt.scatter(data[:,0], data[:,1], c=colors)
```

<!-- grade_id: cell-5c2d005f666d2961, points: 0, grade: true, solution: false -->
```python
correct_w = np.array([ 1.        ,  0.51003169, -1.56964452])
np.testing.assert_allclose(logistic_regression(data, labels, squared_error, squared_error_derivative, 0.5),
                           correct_w,
                           rtol=1e-5)
```

## Problem 6

Implement a single perceptron below.

You can assume that both **x** and **w** will be iterables and b will be a single value. All of them will not be *None*.

<!-- grade_id: cell-6ad288911fa0e2a9, points: 0, grade: false, solution: true -->
```python
def perceptron(x, w, b):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-0beeecf3cc6d7e32, points: 1, grade: true, solution: false -->
```python
assert perceptron([1, 2], [1, 1], 0) == 1
assert perceptron([1, 2], [1, 1], -5) == 0
```